# Phase 2 Code Challenge

This code challenge is designed to test your understanding of the Phase 2 material. It covers:

- Normal Distribution
- Statistical Tests
- Bayesian Statistics
- Linear Regression

_Read the instructions carefully_. You will be asked both to write code and to answer short answer questions.

## Code Tests

We have provided some code tests for you to run to check that your work meets the item specifications. Passing these tests does not necessarily mean that you have gotten the item correct - there are additional hidden tests. However, if any of the tests do not pass, this tells you that your code is incorrect and needs changes to meet the specification. To determine what the issue is, read the comments in the code test cells, the error message you receive, and the item instructions.

## Short Answer Questions 

For the short answer questions...

* _Use your own words_. It is OK to refer to outside resources when crafting your response, but _do not copy text from another source_.

* _Communicate clearly_. We are not grading your writing skills, but you can only receive full credit if your teacher is able to fully understand your response. 

* _Be concise_. You should be able to answer most short answer questions in a sentence or two. Writing unnecessarily long answers increases the risk of you being unclear or saying something incorrect.

```python
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = 'all'
```

```python
# Run this cell without changes to import the necessary libraries

import itertools
import numpy as np
import pandas as pd 
from numbers import Number
from scipy import stats
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline

import pickle
```

---
## Part 1: Normal Distribution [Suggested time: 20 minutes]
---
In this part, you will analyze check totals at a TexMex restaurant. We know that the population distribution of check totals for the TexMex restaurant is normally distributed with a mean of \\$20 and a standard deviation of \\$3. 

### 1.1) Create a numeric variable `z_score_26` containing the z-score for a \\$26 check. 

**Starter Code**

    z_score_26 = 

<!-- grade_id: cell-b947773fd0a27e5d, points: 0, grade: false, solution: true -->
```python
# your code here
mu = 20
std = 3
x = 26
z_score_26 = (x-mu)/std

z_score_26
```

Output:
```
2.0
```

<!-- grade_id: cell-a72f04fffcbada78, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named z_score_26

assert isinstance(z_score_26, Number)

```

### 1.2) Create a numeric variable `p_under_26` containing the approximate proportion of all checks that are less than \\$26.

Hint: Use the answer from the previous question along with the empirical rule, a Python function, or this [z-table](https://www.math.arizona.edu/~rsims/ma464/standardnormaltable.pdf).

**Starter Code**

    p_under_26 = 

<!-- grade_id: cell-a223448d6f75a998, points: 0, grade: false, solution: true -->
```python
# your code here
p_under_26 = stats.norm.cdf(z_score_26)
p_under_26
```

Output:
```
0.9772498680518208
```

<!-- grade_id: cell-626d2e6dea5a86ac, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named p_under_26

assert isinstance(p_under_26, Number)

# These tests confirm that p_under_26 is a value between 0 and 1

assert p_under_26 >= 0
assert p_under_26 <= 1

```

### 1.3) Create numeric variables `conf_low` and `conf_high` containing the lower and upper bounds (respectively) of a 95% confidence interval for the mean of one waiter's check amounts using the information below. 

One week, a waiter gets 100 checks with a mean of \\$19 and a standard deviation of \\$3.

**Starter Code**

    n = 100
    mean = 19
    std = 3
    
    conf_low = 
    conf_high = 

<!-- grade_id: cell-fb0a963f092cb64d, points: 0, grade: false, solution: true -->
```python
# your code here
n = 100
mean = 19
std = 3
c= .95
confidence_interval = stats.norm.interval(c, loc = mean, scale = std/np.sqrt(n))
a,b = confidence_interval
conf_low = a
conf_high = b

```

```python
a
```

Output:
```
18.412010804637983
```

```python
b
```

Output:
```
19.587989195362017
```

```python
a,b = confidence_interval
```

<!-- grade_id: cell-d995ba92d5f4f1f3, points: 1, grade: true, solution: false -->
```python
# These tests confirm that you have created numeric variables named conf_low and conf_high

assert isinstance(conf_low, Number)
assert isinstance(conf_high, Number)

# This test confirms that conf_low is below conf_high

assert conf_low < conf_high

# These statements print your answers for reference to help answer the next question

print('The lower bound of the 95% confidence interval is {}'.format(conf_low))
print('The upper bound of the 95% confidence interval is {}'.format(conf_high))

```

Output:
```
The lower bound of the 95% confidence interval is 18.412010804637983
The upper bound of the 95% confidence interval is 19.587989195362017
```

### 1.4) Short Answer: Interpret the 95% confidence interval you just calculated in Question 1.3.

<!-- grade_id: cell-8212712ac7485fa9, points: 1, grade: true, solution: true -->
95% confident that the true mean of the waiter's checks is between the lower bound and upper bound.

---
## Part 2: Statistical Testing [Suggested time: 20 minutes]
---
The TexMex restaurant recently introduced queso to its menu.

We have a random sample containing 2000 check totals, all from different customers: 1000 check totals for orders without queso ("no queso") and 1000 check totals for orders with queso ("queso").

In the cell below, we load the sample data for you into the arrays `no_queso` and `queso` for the "no queso" and "queso" order check totals, respectively.

```python
# Run this cell without changes

# Load the sample data 
no_queso = pickle.load(open('./no_queso.pkl', 'rb'))
queso = pickle.load(open('./queso.pkl', 'rb'))
```

Output:
```
FileNotFoundError                         Traceback (most recent call last)
<ipython-input-12-96366659c885> in <module>
      2 
      3 # Load the sample data
----> 4 no_queso = pickle.load(open('./no_queso.pkl', 'rb'))
      5 queso = pickle.load(open('./queso.pkl', 'rb'))

FileNotFoundError: [Errno 2] No such file or directory: './no_queso.pkl'
```

### 2.1) Short Answer: State null and alternative hypotheses to use for testing whether customers who order queso spend different amounts of money from customers who do not order queso.

<!-- grade_id: cell-1992605f4c2b8751, points: 1, grade: true, solution: true -->
Null = Customers who order queso do not spend different amounts of money from customers who do not order queso.
Alt = Customers who order queso do spend different amounts of money from customers who do not order queso

### 2.2) Short Answer: What would it mean to make a Type I error for this specific hypothesis test?

Your answer should be _specific to this context,_  not a general statement of what Type I error is.

<!-- grade_id: cell-f5e449677777b2fb, points: 1, grade: true, solution: true -->
Type 1 error = Rejected the null when in reality the data shows that customers who order queso do not spend different amounts of money from customers who do not order queso

Type 2 error = Did not reject the null when in reality the data shows that customers who order queso do spend different amounts of money from customers who do not order queso.

### 2.3) Create a numeric variable `p_value` containing the p-value associated with a statistical test of your hypotheses. 

You must identify and implement the correct statistical test for this scenario. You can assume the two samples have equal variances.

Hint: Use `scipy.stats` to calculate the answer - it has already been imported as `stats`. Relevant documentation can be found [here](https://docs.scipy.org/doc/scipy/reference/stats.html#statistical-tests).

**Starter Code**

    p_value = 

<!-- grade_id: cell-c2e90acfaebc8303, points: 0, grade: false, solution: true -->
```python
# your code here
t,p_value = stats.ttest_ind(queso, no_queso)
```

```python
p_value
```

<!-- grade_id: cell-2364f38312cbc15e, points: 1, grade: true, solution: false -->
```python
# These tests confirm that you have created a numeric variable named p_value

assert isinstance(p_value, Number)

```

### 2.4) Short Answer: Can you reject the null hypothesis using a significance level of $\alpha$ = 0.05? Explain why or why not.

<!-- grade_id: cell-92ea8835a9a326eb, points: 1, grade: true, solution: true -->
Based on the alpha of .05 and the low p value we find that we can reject the null and that the alt hypothesis is statistically significant.

---
## Part 3: Bayesian Statistics [Suggested time: 15 minutes]
---
A medical test is designed to diagnose a certain disease. The test has a false positive rate of 10%, meaning that 10% of people without the disease will get a positive test result. The test has a false negative rate of 2%, meaning that 2% of people with the disease will get a negative result. Only 1% of the population has this disease.

### 3.1) Create a numeric variable `p_pos_test` containing the probability of a person receiving a positive test result.

Assume that the person being tested is randomly selected from the broader population.

**Starter Code**
    
    false_pos_rate = 0.1
    false_neg_rate = 0.02
    population_rate = 0.01
    
    p_pos_test = 

<!-- grade_id: cell-30ff893c7c30f29d, points: 0, grade: false, solution: true -->
```python
# your code here
false_pos_rate = 0.1
false_neg_rate = 0.02
population_rate = 0.01
# probability of sick & healthy
p_sick = .01
p_not_sick = 1 - .01

# probability of test being correct (accuracy)
p_positive_sick = 1 - .1
p_positive_not_sick = 1-.90

# probability of positive test (whether or not you are sick)
p_pos_test = p_sick*p_positive_sick + p_not_sick*p_positive_not_sick

p_pos_test
```

Output:
```
0.10799999999999998
```

<!-- grade_id: cell-2d49a1068c9e46b2, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named p_pos_test

assert isinstance(p_pos_test, Number)

# These tests confirm that p_pos_test is a value between 0 and 1

assert p_pos_test >= 0
assert p_pos_test <= 1

```

### 3.2) Create a numeric variable `p_pos_pos` containing the probability of a person actually having the disease if they receive a positive test result.

Assume that the person being tested is randomly selected from the broader population.

Hint: Use your answer to the previous question to help answer this one.

**Starter Code**
    
    false_pos_rate = 0.1
    false_neg_rate = 0.02
    population_rate = 0.01
    
    p_pos_pos = 

<!-- grade_id: cell-742273642f14d1ab, points: 0, grade: false, solution: true -->
```python
# your code here
false_pos_rate = 0.1
false_neg_rate = 0.02
population_rate = 0.01

sick_prior = .01

# how likely are we to test positive and be sick
pos_sick_likelihood = .9 

# Add normalization
pos_norm = p_pos_test

#finish the find_posterior function 
def find_posterior(prior, likelihood, norm):
    return (likelihood * prior)/ norm


p_pos_pos = find_posterior(sick_prior, pos_sick_likelihood, pos_norm)
p_pos_pos
```

Output:
```
0.08333333333333333
```

<!-- grade_id: cell-4825f55ab11b5fff, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named p_pos_pos

assert isinstance(p_pos_pos, Number)

# These tests confirm that p_pos_pos is a value between 0 and 1

assert p_pos_pos >= 0
assert p_pos_pos <= 1

```

---

## Part 4: Linear Regression [Suggested Time: 20 min]
---

In this section, you'll run regression models with advertising data. The dataset includes the advertising spending and sales for 200 products. These products are a random sample of 200 products from a larger population of products. 

The dataset has three features - `TV`, `radio`, and `newspaper` - that describe how many thousands of advertising dollars were spent promoting the product. The target, `sales`, describes how many millions of dollars in sales the product had.

```python
# Run this cell without changes

import statsmodels
from statsmodels.formula.api import ols

data = pd.read_csv('./advertising.csv').drop('Unnamed: 0', axis=1)
X = data.drop('sales', axis=1)
y = data['sales']
```

### 4.1) Create a variable `tv_mod_summ` containing the summary of a fit StatsModels `ols` model for a linear regression using `TV` to predict `sales`. 

**Starter Code**

    tv_mod = ols(
    tv_mod_summ = tv_mod.summary()

<!-- grade_id: cell-86323454c9b2aa75, points: 0, grade: false, solution: true -->
```python
# your code here
simple_formula = 'sales ~ TV'


tv_mod = ols(simple_formula, data).fit()
tv_mod_summ = tv_mod.summary()
```

<!-- grade_id: cell-84e163360346b9f6, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a variable named tv_mod_summ containing a StatsModels Summary

assert type(tv_mod_summ) == statsmodels.iolib.summary.Summary

# This statement prints your answer for reference to help answer the next question

print(tv_mod_summ)

```

### 4.2) Short Answer: Is there a statistically significant relationship between TV advertising spend and sales in this model? How did you determine this from the model output? 

This question is asking you to use your findings from the sample in your dataset to make an inference about the relationship between TV advertising spend and sales in the broader population.

<!-- grade_id: cell-ebd4cb97464e0212, points: 1, grade: true, solution: true -->
It is statistically significant because the R2 is .612 and the p value of TV ads is 0.
Look at coef of TV (0.0475). When X(TV) increases by 1, the sales increases on average by 0.0475 sales.

### 4.3) Short Answer: Run the cell below to produce a correlation matrix. Given the output, would you expect any collinearity issues if you included all of these features in one regression model? 

```python
# Run this cell

X.corr()
```

<!-- grade_id: cell-c1e98a26b16748b1, points: 1, grade: true, solution: true -->


### 4.4) Create a variable `all_mod_summ` containing the summary of a fit StatsModels `ols` model for a multiple regression using `TV`, `radio`, and `newspaper` to predict `sales`. 

**Starter Code**

    all_mod = ols(
    all_mod_summ = all_mod.summary()

<!-- grade_id: cell-a9cc84d7d7666b0f, points: 0, grade: false, solution: true -->
```python
# your code here
simple_formula_all = 'sales ~ TV + radio + newspaper'
all_mod = ols(simple_formula_all, data).fit()
all_mod_summ = all_mod.summary()


```

<!-- grade_id: cell-1f7f745c780b07b6, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a variable named all_mod_summ containing a StatsModels Summary

assert type(all_mod_summ) == statsmodels.iolib.summary.Summary

# This statement prints your answer for reference to help answer the next question

print(all_mod_summ)

```

### 4.5) Short Answer: Does this model do a better job of explaining sales than the previous model using only the `TV` feature? Explain how you determined this based on the model output. 

<!-- grade_id: cell-7e27f10fcaaed6a9, points: 1, grade: true, solution: true -->
This model does do a better job explaining the sales column because the R2 is much higher due to the significance of the Radio advertisments weighing heavily on the amount of sales. We can see this by looking at the coef radio (0.1885). When X(radio) increases by 1, the sales increases on average by 0.1885 sales.

The newspaper column has no affect so it could be dropped but it would not change the outcome significantly in terms of sales.

```python

```